<a href="https://colab.research.google.com/github/isabellelam/dsc207r/blob/main/week1.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

# Week 1 Study Notebook

# UC San Diego OMDS DSC 207

## Check your Python skills while using Jupyter notebooks.

```python

```

### Installing Otter-Grader

In this course, we will be using Otter-Grader to evaluate your code using predefined test cases. Otter-Grader is a powerful tool that helps you check your work before submitting assignments. To get started with Otter-Grader, you need to install it in your Jupyter Notebook environment.

To install Otter-Grader, simply run the below code cell.

```python
# DO NOT MODIFY
%pip install otter-grader

files = "https://github.com/dsc207rfall2023/fa23course-dsc207r/raw/main/assignments/Week%201/tests.zip"
!wget $files && unzip -o tests.zip
```

```python
# Initialize Otter
import otter
grader = otter.Notebook()
```

### Loading Python modules

```python
import pandas as pd
import numpy as np

# ADD ADDITIONAL PYTHON MODULES BELOW
```

### 1.  Write a program to swap a string and an int variable and print their data types to verify the swapping operation

* Initialize 2 variables my_int = 5 , and my_str = "Hi"
* Swap these variables such that my_int now holds the value of my_float and vice versa
* Verify this by printing their data types

_Points:_ 5

```python

my_int = ...
my_str = ...
```

```python
grader.check("q1")
```

### 2. Write a Python Program to Calculate the length of the diagonal of a rectangle

* Initialize the 2 sides of the rectangle a and b to 5 and 6
* Calculate the length of the diagonal using the formula - diagonal = √(a^2 + b^2)

_Points:_ 5

```python
def calc_diag(a, b):
    ...
```

```python
grader.check("q2")
```

### 3. Write a Python Program to Check if a Number is divisible by 3

* Initialize a variable that takes input from user
* Check and return True or False if the number is divisible by 3

_Points:_ 10

```python
def is_divisible_by_3(num):
    if not isinstance(num, int):
        return False

    if num % 3 == 0:
        return True
    else:
        return False
```

```python
grader.check("q3")
```

<!-- BEGIN QUESTION -->

### 4. Write a Python Program to Check if a number has exactly 3 digits(no zeros allowed at the beginning and the number can be assumed to be positive)

* Initialize a variable that takes input from user(input will be given as a string)
* Check and print if the number has exactly 3 digits(no zeros allowed at the beginning)
* Read about type casting in python and use type casting from string to int

_Points:_ 10

```python

num = ...
```

<!-- END QUESTION -->

### 5. Write a Python Program to Check Whether a String is Palindrome or Not

* Initialize a string
* Lower case it to allow easy comparison
* Check using if-else statements if the string is a palindrome
* A palindrome is a string which is same read forward or backwards.

_Points:_ 10

```python
def if_palindrome(x):
    ...
```

```python
grader.check("q5")
```

### 6.  **Print numbers divisible by 5 between 4 and 30, inclusive**

* Iterate through the numbers in the for loop and notice the way range has been used to generate numbers in the for loop
* Return a list containing only the numbers divisible by 5 using if statement

_Points:_ 5

```python

sol = ...
```

```python
grader.check("q6")
```

### 7. Write a Python Program to print the first n numbers of the Fibonacci series starting from 0

* Initialize a variable n that denotes the number of elements of the Fibonacci series to be computed
* Use a loop and return a list of first n numbers of the Fibonacci series: 0,1,1,2,3,....


_Points:_ 15

```python
def fibonacci(n):
    ...

```

```python
grader.check("q7")
```

### 8. Write a Python Program that defines a function to Remove Vowels From a String

* You can initialize a string of your choice
* Define the function
* Use a for loop to remove vowels

_Points:_ 10

```python
def remove_vowels(input_str):
    ...

```

```python
grader.check("q8")
```

### 9. **Write a program to define a function that takes two numbers as arguments and returns the sum of all numbers between the arguments passed (exclusive, i.e. excluding the limits)**

* You can define a range of your choice
* Define the function
* Use a for loop to add numbers within the range(exclusive)

_Points:_ 10

```python
def sum(a, b):
    ...
```

```python
grader.check("q9")
```

<!-- BEGIN QUESTION -->

### 10. Define a class named Triangle used to calculate the area

* The class can be constructed by the length of the sides
* The class must have a method which can compute the area using the formula area = √(s(s-a)(s-b)(s-c)) where s=(a+b+c)/2 is the semi-perimeter
* The class must verify if the side lengths can form a triangle

_Points:_ 20

```python
class Triangle:
    ...

    def __init__():
        ...

    def valid_triangle():
        ...

    def calculate_area():
        ...


triangle1 = Triangle(4, 3, 2)
area1 = triangle1.get_area()
print(f'The area of triangle1 is: {area1}')

triangle2 = Triangle(2, 4, 6)
area2 = triangle2.get_area()
print(f'The area of the triangle2 is: {area2}')

triangle3 = Triangle(7, 9, 12)
area3 = triangle3.get_area()
print(f'The area of the triangle3 is: {area3}')
```

<!-- END QUESTION -->

### Submission

Before you submit your notebook, ensure that you've run all the cells sequentially to display images, graphs, and outputs correctly. Take the time to review your solutions and evaluate using the public test cases. Once you're satisfied, save your notebook by navigating to "File" in the Jupyter Notebook menu and selecting "Save" or using the keyboard shortcut (usually Ctrl + S or Cmd + S).

To submit, export your notebook as an .ipynb file. Ensure it's saved with the same name as the assignment. Then, upload this .ipynb file to Gradescope. Remember to submit the correct version.

Please save your work before exporting it, and if you encounter any technical issues or have questions about the submission process, reach out to the course staff for assistance

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

Tests for q1 (points: 5)
```
>>> assert type(my_int) == str
```
```
>>> assert type(my_str) == int
```
```
>>> assert my_int == "Hi"
```
```
>>> assert my_str == 5
```


Tests for q2 (points: 5)
```
>>> assert calc_diag(3, 4) == 5
```


Tests for q3 (points: 10)
```
>>> assert is_divisible_by_3(9) == True
```
```
>>> assert is_divisible_by_3(-12) == True
```


Tests for q5 (points: 10)
```
>>> assert if_palindrome("racecar") == True
```
```
>>> assert if_palindrome("hello") == False
```


Tests for q6 (points: 5)
```
>>> assert len(sol) == 6
```
```
>>> assert sol == [5, 10, 15, 20, 25, 30]
>>>
```


Tests for q7 (points: 15)
```
>>> assert fibonacci(1) == [0]
```
```
>>> assert fibonacci(5) == [0, 1, 1, 2, 3]
```


Tests for q8 (points: 10)
```
>>> assert remove_vowels("Hello, World!") == "Hll, Wrld!"
```
```
>>> assert remove_vowels("Python Programming") == "Pythn Prgrmmng"
```


Tests for q9 (points: 10)
```
>>> assert sum(1, 5) == 9
>>>
```
```
>>> assert sum(-3, 3) == 0
```
